# RO47002 Machine Learning for Robotics
* (c) TU Delft, 2020
* Period: 2020-2021, Q1
* Course homepage: https://brightspace.tudelft.nl/d2l/home/318952

Make sure you fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Moreover, if you see an empty cell, please DO NOT delete it, instead run that cell as you would run all other cells. Please fill in your name(s) and other required details below:

```python
GROUP_NUMBER = ""
STUDENT_NAME1 = ""
STUDENT_NUMBER1 = ""
STUDENT_NAME2 = ""
STUDENT_NUMBER2 = ""
```

<!-- grade_id: cell-c35e4c8223095209, points: 0, grade: true, solution: false -->
```python
# Note: this block is a check that you have filled in the above information.
# It will throw an AssertionError until all fields are filled
assert(GROUP_NUMBER != "")
assert(STUDENT_NAME1 != "")
assert(STUDENT_NUMBER1 != "")
assert(STUDENT_NAME2 != "")
assert(STUDENT_NUMBER2 != "")
```

### General announcements

* Do *not* share your solutions, and do *not* copy solutions from others. By submitting your solutions, you claim that you and your lab partner alone are responsible for this code.

* Do *not* email questions directly, since we want to provide everybody with the same information and avoid repeating the same answers. Instead, please post your questions regarding this assignment in the correct support forum on Brightspace, this way everybody can benefit from the response. If you do have a particular question that you want to ask directly, please use the scheduled practicum hours to ask a TA.

* There is a strict deadline for each assignment. Students are responsible to ensure that they have uploaded their work in time. So, please double check that your upload succeeded to the Brightspace and avoid any late penalties.

* This [Jupyter notebook](https://jupyter.org/) uses `nbgrader` to help us with automated tests. `nbgrader` will make various cells in this notebook "uneditable" or "unremovable" and gives them a special id in the cell metadata. This way, when we run our checks, the system will check the existence of the cell ids and verify the number of points and which checks must be run. While there are ways that you can edit the metadata and work around the restrictions to delete or modify these special cells, you should not do that since then our nbgrader backend will not be able to parse your notebook and give you points for the assignment. You are free to add additional cells, but if you find a cell that you cannot modify or remove, please know that this is on purpose.

* This notebook will have in various places a line that throws a `NotImplementedError` exception. These are locations where the assignment requires you to adapt the code! These lines are just there as a reminder for youthat you have not yet adapted that particular piece of code, especially when you execute all the cells. Once your solution code replaced these lines, it should accordingly *not* throw any exceptions anymore.

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

---

# week 5

* Topic: Dimensionality reduction, clustering
* Year: 2020-2021
* Book chapters: 8, 9

## Introduction

This lab assignment consists of two parts.

**This is Part 2 - Clustering**

In this part you will:

* Using K-Means on a toy dataset
* Implementing the basic K-Means algorithm yourself
* Using a Gaussian Mixture Model (GMM) for clustering
* Comparing K-Means and GMM for outlier detection
* Application: *Meeting and recognizing human faces*
* A simple semi-supervised approach to improve clustering

```python
# Python ≥3.5 is required
import sys
assert sys.version_info >= (3, 5)

# Scikit-Learn ≥0.20 is required
import sklearn
assert sklearn.__version__ >= "0.20"

# Common imports
import numpy as np
import os

# to make this notebook's output stable across runs
np.random.seed(42)

# To plot pretty figures
#%matplotlib inline
import matplotlib as mpl
import matplotlib.pyplot as plt
mpl.rc('axes', labelsize=14)
mpl.rc('xtick', labelsize=12)
mpl.rc('ytick', labelsize=12)
```

```python
import ipywidgets
from IPython.display import display
```

# Clustering

We start again by exploring some clustering techinques on a toy dataset,
which is generated here:

```python
from sklearn.datasets import make_blobs

# Create toy dataset
n_samples = 1000

X, y = make_blobs(n_samples=n_samples,
    cluster_std=[0.5, 2.0, 1.5],
    random_state=15)
```

```python
# Show plot of the data
def plot_data_with_labels(X, y=None):
    s = plt.scatter(X[:,0], X[:,1], c=y)
    if not y is None:
        plt.legend(*s.legend_elements())
        
plot_data_with_labels(X)
```

One common algorithm used to cluster data is **k-means**.
In the block below,

* use sklearn's k-Means implementation to cluster the given data `X`, and predict each sample's cluster label. Initialize k-Means by assuming that there are $k=3$ clusters in the data.
* also obtain the predictions of the fitted k-Means solutions on all samples in `X`, and call this `y_pred_km`.

<!-- grade_id: cell-c858924165ff1e40, points: 0, grade: false, solution: true -->
```python
k = 3

kmeans = None # store sklearns k-Means implementation here
y_pred_km = None # store the predicted cluster labels for data X here 

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-b1698a40756eb52e, points: 1, grade: true, solution: false -->
```python
assert( kmeans != None )
assert( len(np.unique(y_pred_km)) == 3 )
assert( y_pred_km.shape == (n_samples,) )

```

Let's checkout the cluster labels that the algorithm found:

```python
# color the samples with the predicted labels
plot_data_with_labels(X, y_pred_km)

# also show the cluster centers found by k-Means
plt.plot(kmeans.cluster_centers_[:,0], kmeans.cluster_centers_[:,1], 'rd')
```

Alternatively, we could fit a Gaussian Mixture distribution on the data,
and label the samples based on their probability under the fitted mixture components.

In the code below, fit a Gaussian Mixture on the data using $k=3$ mixture components.

<!-- grade_id: cell-3c71cdcaf943a55c, points: 0, grade: false, solution: true -->
```python
gm = None # store your sklearn GaussianMixture here
y_pred_gm = None # store the predicted class labels from the GaussianMixture here

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e1e3b25f849405be, points: 1, grade: true, solution: false -->
```python
assert( gm != None )
assert( len(np.unique(y_pred_gm)) == 3 )
assert( y_pred_gm.shape == (n_samples,) )

```

Let's again visualize the result, and compare it to k-Means.
You should see that the plots look pretty similar, except perhaps for a few samples near the cluster boundaries.

```python
plt.figure(figsize=(12,4))
plt.subplot(1,2,1)
plt.plot(kmeans.cluster_centers_[:,0], kmeans.cluster_centers_[:,1], 'rd')
plot_data_with_labels(X, y_pred_km)
plt.title('k-Means')

plt.subplot(1,2,2)
plt.plot(gm.means_[:,0], gm.means_[:,1], 'rd')
plot_data_with_labels(X, y_pred_gm)
plt.title('Gaussian Mixture')
```

## Implement k-means algorithm

The basic k-Means algorithm actually consists only a few steps:

1. initialize k cluster centers by picking k random samples
2. assign samples to nearest cluster center
3. for each cluster k, compute the new center as the mean of the assigned samples
4. goto step 2, until convergence or sufficient iterations have passed

In the code blocks below, you will implement steps 1 to 3.

Implement **step 1**: select k samples as random samples from X and set these as cluster centers.
*HINT:* you can use `np.random.choice` to randomly select k indices from a given range.

<!-- grade_id: cell-b8cd2167361538be, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-15ba7d96be57b058, points: 2, grade: true, solution: false -->
```python
k = 3

np.random.seed(5)
centers = initialize_centers(k, X)

assert(centers.shape == (k, 2)) # k centers in a 2D feature space

centers_init = centers.copy() # keep a copy to compare later to

# let's plot the randomly selected cluster centers
plot_data_with_labels(X)
plt.plot(centers[:,0], centers[:,1], 'rd')
plt.title('Initial cluster centers (no data assigned yet)')
```

Implement **step 2**: assign all data points to nearest centers (in terms of Euclidean distance)
*Hint:* you can use `scipy.spatial.distance.cdist` to quickly compute the distance of the samples to the cluster centers.


<!-- grade_id: cell-08c553fc7f5d3ec4, points: 0, grade: false, solution: true -->
```python
import scipy

def assign_samples_to_centers(centers, X):
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return cluster_idxs
```

<!-- grade_id: cell-167bfcea73922c65, points: 2, grade: true, solution: false -->
```python
cluster_idxs = assign_samples_to_centers(centers, X)

assert( cluster_idxs.shape == (n_samples,) )
assert(cluster_idxs.min() >= 0) # lowest possible cluster id is 0
assert(cluster_idxs.max() < k) # highest possible cluster id is k-1

# Let's see how the samples were assigned to the randomly picked cluster centers
plt.plot(centers[:,0], centers[:,1], 'rd')
plot_data_with_labels(X, cluster_idxs)
plt.title('Assigning data to nearest cluster centers')
```

Implement **step 3**: compute new center means based on the sasigned cluster labels.
In this step, iterate over the k cluster labels, and compute the center of cluster $c$ as the mean of all samples in X assigned to $c$ in `cluster_idxs`.
The result is a new set of $k$ cluster centers, based on the current assignment.

<!-- grade_id: cell-2595920badf8056e, points: 0, grade: false, solution: true -->
```python
def update_center_means(k, X, cluster_idxs):
    # YOUR CODE HERE
    raise NotImplementedError()
    return centers
```

<!-- grade_id: cell-84d75b6d06f813a9, points: 2, grade: true, solution: false -->
```python
centers = update_center_means(k, X, cluster_idxs)

assert(centers.shape == (k, 2)) # k centers in a 2D feature space

# check that the centers have changed w.r.t. the initial centers that we stored in center_init above
assert(np.all( centers_init != centers ))

# Show the new assignment
plt.plot(centers[:,0], centers[:,1], 'rd')
plot_data_with_labels(X, cluster_idxs)
plt.title('After recomputing the cluster centers')
```

That's it! If you have implemented these steps correctly, we can run all step sequentially for a few iterations.
Notice how after a few iterations already the cluster centers are not moving much anymore.
Once we detect that the number of samples assigned to each cluster haven't changed, we can quit iterating.

```python
k = 3

# Step 1
centers = initialize_centers(k, X)

plt.plot(centers[:,0], centers[:,1], 'rd')
plot_data_with_labels(X)
plt.title('initialization')
plt.show()

max_iters = 10
last_cluster_counts = np.zeros(k)
for j in range(max_iters):
    # Step 2
    cluster_idxs = assign_samples_to_centers(centers, X)
    
    # Step 3
    centers = update_center_means(k, X, cluster_idxs)

    # Show intermediate result
    cluster_counts = np.bincount(cluster_idxs, minlength=k)
    plt.title(f'iter {j}: ' + str(cluster_counts))
    plt.plot(centers[:,0], centers[:,1], 'rd')
    plot_data_with_labels(X, cluster_idxs)
    plt.show()
    
    # Step 4: continue with next iteration, unless converged
    if np.all(cluster_counts == last_cluster_counts):
        # detected cluster assignment didn't change anymore since last iteration,
        # so cluster centers have converged now
        print(f'Converged in {j} iterations!')
        break # quit the for-loop
    else:
        # keep this iteration's cluster counters to compare to in next iteration
        last_cluster_counts = cluster_counts
```

## Finding the outliers

Let's try to do some anomaly detection, that is finding outliers which are dissimilar to most of the rest of the taining data. One approach one might come up with is to consider any sample "sufficiently" far away from a cluster center as an outlier. With k-Means we can reuse the scoring metric used to assign samples to cluster centers, thus a sample is considered an outlier if its Euclidean distance to nearest cluster center is above a certain distance threshold.

Complete the function `find_kmeans_distance_outliers()` below to implement this outlier detection method.

<!-- grade_id: cell-21916fcd16d7f444, points: 0, grade: false, solution: true -->
```python
def find_kmeans_distance_outliers(kmeans, X, dist_thresh):
    """
    Determine for each sample in X if it is an outlier or not.
    Outliers are found by finding the distance to the closest cluster center.
    If this smallest distance is above the given threshold `dist_thresh`,
    then the sample is an outlier
    
    Input: kmeans - a fitted instance of sklearn's KMeans class
    Input: X - a N x M dataset of N samples with M features
    Input: dist_thresh - a number, outliers are more distance than this to all cluster centers
    Returns: is_outlier - a N-dimensional Boolean numpy vector,
             is_outlier[i] is True only if sample X[i] is an outlier.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return is_outlier
```

<!-- grade_id: cell-3b424ecaeb1ad495, points: 3, grade: true, solution: false -->
```python
is_outlier = find_kmeans_distance_outliers(kmeans, X, dist_thresh=2.0)
assert(len(is_outlier) == n_samples)
assert(is_outlier.dtype == np.bool)

assert(is_outlier[0] == False) # this sample should not be an outlier for this threshold
assert(is_outlier[2] == True) # this sample should be an outlier for this threshold
assert(is_outlier[3] == True) # this sample should be an outlier for this threshold
assert((sum(is_outlier) > 320) and (sum(is_outlier) < 350)) # total number of expected outliers ~332

is_outlier = find_kmeans_distance_outliers(kmeans, X, dist_thresh=5.0)
assert(is_outlier[0] == False) # this sample should not be an outlier for this threshold
assert(is_outlier[2] == True) # this sample should be an outlier for this threshold
assert(is_outlier[3] == False) # this sample should not be an outlier for this threshold
assert((sum(is_outlier) > 2) and (sum(is_outlier) < 8)) # total number of expected outliers ~5
```

```python
# Create interactive widget to explore outliers
def show_kmeans_outliers(dist_thresh):
    is_outlier = find_kmeans_distance_outliers(kmeans, X, dist_thresh)
    
    plt.plot(X[~is_outlier,0], X[~is_outlier,1], '.', label='inlier')
    plt.plot(X[is_outlier,0], X[is_outlier,1], '.', label='outlier')
    plt.plot(centers[:,0], centers[:,1], 'rd')
    plt.legend()
    plt.grid('on')

ipywidgets.interact(show_kmeans_outliers, dist_thresh=(0., 5.0))
```

However, a better way could be to use a density estimation technique, such as a Gaussian Mixture Model, as seen in the book. To use the Gaussian Mixture for outlier detection, we could now score the samples with respect to the distribution fit on all the data,  i.e. we compute the log probability density of each sample under the distribution.
An outlier is than a sample for which the log probability is under a given threshold

As you did for k-Means outlier detection, complete the function below:

<!-- grade_id: cell-c90fa1f7675f118f, points: 0, grade: false, solution: true -->
```python
def find_gm_logporb_outliers(gm, X, logprob_thresh):
    """
    Determine for each sample in X if it is an outlier or not.
    Outliers are found by finding the log-probability under the fitted Gaussian mixture.
    If this log probability is BELOW the given threshold `logprob_thresh`,
    then the sample is an outlier.
    
    Input: gm - a fitted instance of sklearn's GaussianMixture class
    Input: X - a N x M dataset of N samples with M features
    Input: logprob_thresh - a number, outliers how a log-probability under the GM below this threshold
    Returns: is_outlier - a N-dimensional Boolean numpy vector,
             is_outlier[i] is True only if sample X[i] is an outlier.
    """

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return is_outlier
```

<!-- grade_id: cell-ba29f37e6837d6f7, points: 2, grade: true, solution: false -->
```python
is_outlier = find_gm_logporb_outliers(gm, X, logprob_thresh=-7.0)
assert(len(is_outlier) == n_samples)
assert(is_outlier.dtype == np.bool)

assert(is_outlier[0] == False) # this sample should not be an outlier for this threshold
assert(is_outlier[2] == True) # this sample should be an outlier for this threshold
assert(is_outlier[3] == False) # this sample should be an outlier for this threshold
assert((sum(is_outlier) > 22) and (sum(is_outlier) < 32)) # total number of expected outliers ~27
```

Let's take a look at the distribution of log probabilities for our dataset, and see if we can get an idea in what range we should put our cutoff threshold:

```python
plt.hist(gm.score_samples(X));
plt.xlabel('log probability')
plt.ylabel('number of samples')
```

We can see that most samples have a log probability above approximately -7.
A cutoff threshold in the range -8 and -6 is probably good, we could reasonably argue that the samples in the tail of this distribution are outliers. Of course in practice the cutoff would be depend on your task, so this is not a general rule.

We can study the effect of different thresholds in the code block below.

```python
# Create interactive widget to explore outliers
def show_gm_outliers(logprob_thresh):
    is_outlier = find_gm_logporb_outliers(gm, X, logprob_thresh)
    
    plt.plot(X[~is_outlier,0], X[~is_outlier,1], '.', label='inlier')
    plt.plot(X[is_outlier,0], X[is_outlier,1], '.', label='outlier')
    plt.plot(gm.means_[:,0], gm.means_[:,1], 'rd')
    plt.legend()
    plt.grid('on')

ipywidgets.interact(show_gm_outliers, logprob_thresh=(-12., .0))
```

Compare the earlier outlier detection method using the Euclidean distance to the cluster center to this approach of using the log probability under the GMM.

**Q**: What happens when you try to detect possible outliers in the smaller cluster in the bottom-right? Explain what you observe with the Euclidean-distance based detection method, and what you observe with the Gaussian Mixture detection method. What causes these differences?

*Hint*: think about how the distance between the a sample and the center (mean) are used in the formula for the Gaussian distribution.

<!-- grade_id: cell-42c2c9f88f758989, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

We end this part with a note.

An even better way to find outliers with respect to the overall data, would be to fit the distribution on all data EXCEPT the sample we are scoring,
since we our now 'testing' the distribution on a sample in the 'training' data.
This is what the book refers to *novelty detection*, but an outlier could be consider a novelty w.r.t. the rest of the training data.

However, since in these examples we are using few mixture components the overall distributions should not change too much
of we keep the test sample included (bias-variance trade-off) and it is much faster to fit the distribution
once and test all N samples, instead of fitting a distribution N times for all test samples separately.

# Clustering faces

For the final part, let's consider a social robot that moves around a building,
and continously bumps into people. Using a trained object detector,
it can detect and locate faces in its camera image.
Using the bounding box, a square patch around a detected face could be extracted.
Over time, the robot would build a dataset of faces, some belonging to the same people and others not.
Can you help the robot determine which faces probably correspond to the same person?

![Pepper, the social robot (image source: Wikipedia)](https://upload.wikimedia.org/wikipedia/commons/thumb/a/a1/SoftBank_pepper.JPG/330px-SoftBank_pepper.JPG)


For this exercise we will use the `Olivetti faces dataset`, containing 400 data samples, with 10 samples for 40 persons.
Each sample is a $64 \times 64$ gray scale image patch, reshaped into a 4096-dimensional feature vector.
The class labels of this dataset represent the person identitiy, so all samples with the same class label are taken from the same person.
We will pretend that these are the (gray scale) face image patches that our robot has collected.
The dataset actually contains a set of face images taken between April 1992 and 
April 1994 at AT&T Laboratories Cambridge. 


Note that the first time you run the code below, sklearn will automatically downloads the data
archive from AT&T. The result will be cached, so once it is on your computer it will not need to download it again.

```python
# #############################################################################
# Download the data, if not already on disk and load it as numpy arrays

ALLOW_DOWNLOAD = True

#data = sklearn.datasets.fetch_lfw_people(min_faces_per_person=70, resize=0.4, download_if_missing=ALLOW_DOWNLOAD)
data = sklearn.datasets.fetch_olivetti_faces(download_if_missing=ALLOW_DOWNLOAD)

# To see more information on the downloaded dataset, execute the following line:
#print(data['DESCR'])
```

```python
X, y = data['data'], data['target']

print('Dataset size:', X.shape)
print('Unique person ids:', np.unique(y))
```

To get an idea of the contents of this dataset, we can explore the features in this dataset by resizing them to $64 \times 64$ images.

```python
def show_face(idx):
    x = X[idx]
    x = x.reshape((64,64))
    label = y[idx]
    plt.imshow(x, cmap='gray')
    plt.title(f'sample={idx}, label={label}')
    plt.axis('off')

ipywidgets.interact(show_face, idx=(0,399))
```

In the code block below, compare 3 methods of dimensionality reduction to visualize the data in the 4096-dimensional space in a 2D plot, namely:

1. Principal Component Analysis (see Part 1 of the exercises)
2. Linear Discriminant Analysis (see Part 1 of the exercises)
3. t-distributed Stochastic Neighbor Embedding (t-SNE)

Check the sklearn documentation on [t-SNE](https://scikit-learn.org/stable/modules/generated/sklearn.manifold.TSNE.html?highlight=tsne#sklearn.manifold.TSNE) if you are not sure how to use it.

<!-- grade_id: cell-d14a6c7642c6e8b5, points: 0, grade: false, solution: true -->
```python
%%time

# set these variables
X_pca = None # result of dim. reduction with PCA
X_lda = None # result of dim. reduction with LDA
X_tsne = None # result of dim. reduction with t-SNE

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5b4fbed90cd8d36e, points: 2, grade: true, solution: false -->
```python
assert(X_pca.shape == (400, 2))
assert(X_lda.shape == (400, 2))
assert(X_tsne.shape == (400, 2))

```

Afterwards, we can visualize the data in the projected space.
We will also in color the points in all plots using the true class labels, such that we can see if the resulting embedding has placed points from the same class together or not.

```python
def plot_2D_embedded_classes(X, y):
    for c in np.unique(y):
        plt.plot(X[y==c,0], X[y==c,1], '.', label='%d'%c)
    plt.grid()
    plt.axis('equal')

plt.figure(figsize=(14, 4))
plt.subplot(1,3,1)
plot_2D_embedded_classes(X_pca, y)
plt.title('PCA')

plt.subplot(1,3,2)
plot_2D_embedded_classes(X_lda, y)
plt.title('LDA')

plt.subplot(1,3,3)
plot_2D_embedded_classes(X_tsne, y)
plt.title('TSNE')
```

**Q** Which of these 3 dimensionality reduction techniques appear deterministic, and which stochastic (i.e. provide different results every time you run it?)

**Q** Which of these 3 dimensionality reduction techniques are unsupervised, and thus 'blind' to the true class labels (until) we plot them?

**Q** Which of these 3 dimensionality reduction techniques compute non-linear projections?

<!-- grade_id: cell-d68b2299995068a3, points: 0, grade: false, solution: true -->
```python
# To answer, set each variable to either True or False

#Q1
PCA_IS_DETERMINISTIC = None
LDA_IS_DETERMINISTIC = None
TSNE_IS_DETERMINISTIC = None

#Q2
PCA_IS_UNSUPERVISED = None
LDA_IS_UNSUPERVISED = None
TSNE_IS_UNSUPERVISED = None

#Q3
PCA_IS_NONLINEAR_PROJECTION = None
LDA_IS_NONLINEAR_PROJECTION = None
TSNE_IS_NONLINEAR_PROJECTION = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-89f6aaf615f8e477, points: 3, grade: true, solution: false -->
```python
assert(PCA_IS_DETERMINISTIC in (True, False))
assert(LDA_IS_DETERMINISTIC in (True, False))
assert(TSNE_IS_DETERMINISTIC in (True, False))

assert(PCA_IS_UNSUPERVISED in (True, False))
assert(LDA_IS_UNSUPERVISED in (True, False))
assert(TSNE_IS_UNSUPERVISED in (True, False))

assert(PCA_IS_NONLINEAR_PROJECTION in (True, False))
assert(LDA_IS_NONLINEAR_PROJECTION in (True, False))
assert(TSNE_IS_NONLINEAR_PROJECTION in (True, False))

```

## Comparing k-Means and Gaussian Mixture model

Ok, now let's see what happens if we try to recover the 40 identities by clustering this data.

We do need a metric to evaluate success though.
There are many metrics in sklearn to evaluate the cluster quality with respect to some ground truth labels.
Here, we will use `sklearn.metrics.completeness_score()`.

From the [sklearn documentation](https://scikit-learn.org/stable/modules/generated/sklearn.metrics.completeness_score.html#sklearn.metrics.completeness_score):
```
Completeness metric of a cluster labeling given a ground truth.

A clustering result satisfies completeness if all the data points that are members of a given class are elements of the same cluster.

This metric is independent of the absolute values of the labels: a permutation of the class or cluster label values won’t change the score value in any way.
```

Overall, this metric will be score between 0.0 for bad (all the clusters are randomly divided over the true class labels) to 1.0 for good (clusters perfectly align with class labels, up to permutation of which class is which cluster).

Next, try to cluster the data into 40 clusters using

- k-Means, and
- a Gaussian Mixture model.

**WARNING** Gaussian component has a $M \times M$ covariance matrix, which results in a large amount of parameters to be estimated as the number of feature increases (it grows quadratically!).
Fitting a mixture model on such high dimensional feature space can therefore take a *loooooong* time before it converges, and will also be very prone to overfitting!

**Hint:** First apply PCA to reduce the feature space to, say, 50 components before fitting the Gaussian Mixture.

Remember that to debug your code, and you can use `%%time` to figure out how long it takes to run a block on your PC.

<!-- grade_id: cell-1cb6992161033640, points: 0, grade: false, solution: true -->
```python
%%time

kmeans = None # store your KMeans instance in this variable
y_pred_km = None # store your predictions in this variable

k = 40

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d35933b9a3e81c1f, points: 0, grade: false, solution: true -->
```python
%%time

gm = None # store your GaussianMixture in this variable
y_pred_gm = None # store your predictions in this variable

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-62cea27bfac0cd92, points: 3, grade: true, solution: false -->
```python
score_km = sklearn.metrics.completeness_score(y, y_pred_km)
print('         k-Means:', score_km)

score_gm = sklearn.metrics.completeness_score(y, y_pred_gm)
print('Gaussian Mixture:', score_gm)

assert( kmeans != None )
assert( len(np.unique(y_pred_km)) == 40 )
assert( y_pred_km.shape == (400,) )

assert( gm != None )
assert( len(np.unique(y_pred_gm)) == 40 )
assert( y_pred_gm.shape == (400,) )

sklearn.utils.validation.check_is_fitted(kmeans) # will throw error if not fitted
sklearn.utils.validation.check_is_fitted(gm) # will throw error if not fitted

```

Note that you might see some variance in the results everytime you rerun it.

Answer in the block below the following two questions.

* **Q1**: can you conclude which clustering method finds better clusters on this dataset?
* **Q2**: which approach clusters the data more efficiently in terms of processing time, and is this affected by dimensionality reduction?

<!-- grade_id: cell-856cc8fa825cf514, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

For fun, let's visually inspect the clusters found by both methods:

```python
def plot_faces_class(y_pred, c):
    idxs = np.where(y_pred == c)[0]

    plt.figure(figsize=(16,8))
    for j, idx in enumerate(idxs[:15]):
        plt.subplot(3,5,j+1)
        show_face(idx)

print('K-Means')
ipywidgets.interact(lambda cluster: plot_faces_class(y_pred_km, cluster), cluster=(0,40-1))
```

```python
print('Gaussian Mixture')
ipywidgets.interact(lambda cluster: plot_faces_class(y_pred_gm, cluster), cluster=(0,40-1))
```

## Semi-supervised learning

For the final test, let's assume we DO have a few labels, but not all labels.
In fact, let's assume that our robot gets one face with a person id for each person,
but also collects 9 unlablled pictures of each person.
Can we improve our clustering by using the few labelled samples?

This would be a case of *semi-supervised* learning.

```python
# Let's assume that this part of the data is labelled
# every 10th sample is a new person
X_labelled = X[::10,:]
y_labelled = y[::10]
```

As a baseline, let's see how a standard Random Forest classifier would perform when given just this 1-sample per class dataset on the full unlabelled dataset.

**NOTE**: a better way to draw sound conclusions would be to do a proper train-and-test split and hyperparameter tuning, but a Random Forest is a relatively robust classifier with few hyperparameters, so it is sufficient for this demonstration.

```python
import sklearn.ensemble
rf = sklearn.ensemble.RandomForestClassifier()
rf.fit(X_labelled, y_labelled)
y_pred_rf = rf.predict(X)

comp_score = sklearn.metrics.completeness_score(y, y_pred_rf)
acc_score = sklearn.metrics.accuracy_score(y, y_pred_rf)

print('RF completeness:', comp_score)
print('    RF accuracy:', acc_score)

print('Confusion matrix')
plt.matshow(sklearn.metrics.confusion_matrix(y_pred_rf, y))
```

```python
print('Random Forst (supervised on 10 samples)')
ipywidgets.interact(lambda cluster: plot_faces_class(y_pred_rf, cluster), cluster=(0,40-1))
```

The supervised accuracy of the RF is not great ... but it also was given very little data (1 sample per class!)

Let's try to use a GaussianMixture here again, but let's give it a semi-supervised twist:
**use the given labelled samples as initial mean estimates of the 40 Gaussian terms**.

To do this, take care of the following points:

- You will need to use PCA for dimensionality reduction again, try using about 18 components this time
- PCA should be fitted on all data, but you need to project the labelled data to the PCA space again separately to define the means of the GaussianMixture in this reduced space
- To initialize the means of the GaussianMixture with specific values, you need to set the `init_params` parametere to `'random'`, and then use the `means_init` parameter to define the means. See the [sklearn documentation](https://scikit-learn.org/stable/modules/generated/sklearn.mixture.GaussianMixture.html) for more info.

<!-- grade_id: cell-9d42c608723ea944, points: 0, grade: false, solution: true -->
```python
gm2 = None # store your semi-supervised GaussianMixture in this variable
y_pred_gm2 = None # store the predictions of the semi-supervised GaussianMixture in this variable

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-2df81c72fddd435f, points: 4, grade: true, solution: false -->
```python
comp_score_gm2 = sklearn.metrics.completeness_score(y, y_pred_gm2)
acc_score_gm2 = sklearn.metrics.accuracy_score(y, y_pred_gm2)

print('GM semi-supervised completeness:', comp_score_gm2)
print('   GM semi-supervised  accuracy:', acc_score_gm2)

# these scores should be achievable (though there is a bit of randomness ...)
assert(comp_score_gm2 > 0.82)
assert(acc_score_gm2 > 0.75)

assert( gm2 != None )
assert( len(np.unique(y_pred_gm2)) == 40 )
assert( y_pred_gm2.shape == (400,) )

sklearn.utils.validation.check_is_fitted(gm2) # will throw error if not fitted

```

If everything is correct, you should see that semi-supervised training with the GaussianMixture outperforms the fully unsupervised Gaussian Mixture, but also the fully supervised Random Forest!


Another benefit is that the cluster labels now all align with the true class labels, because of the small amount of supervision, as the confusion matrix below will confirm (it should only be diagonal if the cluster labels are equal to the true class labels):

```python
plt.matshow(sklearn.metrics.confusion_matrix(y_pred_gm2, y))
plt.title('GM (semi-supervised)')

plt.matshow(sklearn.metrics.confusion_matrix(y_pred_rf, y))
plt.title('RF (supervised)')
```

```python
print('Semi-supervised Gaussian Mixture')
ipywidgets.interact(lambda cluster: plot_faces_class(y_pred_gm2, cluster), cluster=(0,40-1))
```

This concludes this week's exercises!